Before you turn this lab in, make sure everything runs as expected. **Restart and run all cells** (in the menubar, select Kernel$\rightarrow$Restart & Run All) then check the output is as expected and there are no errors.  Also remember to **Save** before uploading this file (File$\rightarrow$Save and Checkpoint).

---

# <div align="center">Quantum Computer Systems Design</div>&nbsp;  <div align="center">Lab 5</div>

In the videos, you saw the steps to compile and optimize a quantum algorithm.  For this lab, you will again start with a small circuit, the Cuccaro adder.  This by itself is not a quantum algorithm but can be used as a component in algorithms.  The Cuccaro adder takes as input two quantum registers, a carry-in qubit, and a carry-out qubit and as output overwrites one of the registers and the carry-out bit with the sum.  In the previous lab you decomposed adder circuits.  In this lab you will optimize these decomposed circuits to find equivalent circuits with fewer gates.


# 0.1. Cuccaro Adder Definition

Below is the definition of the Cuccaro adder within Qiskit.  Run each of the below cells untill you reach question 1.1 to start.

```python
import numpy as np
import qiskit
from qiskit import QuantumCircuit, ClassicalRegister, QuantumRegister, execute
from qiskit.circuit.quantumregister import Qubit
from qiskit.tools.visualization import circuit_drawer
from qiskit.quantum_info import state_fidelity, random_unitary
from qiskit.quantum_info.operators.predicates import matrix_equal
from qiskit import BasicAer
backend = BasicAer.get_backend('qasm_simulator')
from IPython.display import display
```

```python
# Implementation of the Cuccaro Quantum Adder from
#   https://github.com/jmbaker94/quantumcircuitbenchmarks
# Design from the paper
#   https://arxiv.org/abs/quant-ph/0410184

# Majority gate
maj_c = QuantumCircuit(3, name='MAJ')
maj_c.cx(2, 1)
maj_c.cx(2, 0)
maj_c.ccx(0, 1, 2)
maj = maj_c.to_gate(label='MAJ')
print('MAJ:')
display(maj_c.draw(fold=-1))

# 2-CNOT version of UnMajority and Add gate
uma2_c = QuantumCircuit(3, name='UMA2')
uma2_c.toffoli(0, 1, 2)
uma2_c.cx(2, 0)
uma2_c.cx(0, 1)
uma2 = uma2_c.to_gate(label='UMA2')
print('UMA2:')
display(uma2_c.draw(fold=-1))

# 3-CNOT version of UnMajority and Add gate
# (Allows more parallelism in the circuit)
uma3_c = QuantumCircuit(3, name='UMA3')
uma3_c.x(1)
uma3_c.cx(0, 1)
uma3_c.toffoli(0, 1, 2)
uma3_c.x(1)
uma3_c.cx(2, 0)
uma3_c.cx(2, 1)
uma3 = uma3_c.to_gate(label='UMA3')
print('UMA3:')
display(uma3_c.draw(fold=-1))

def cuccaro_adder(c, cin, a, b, cout, uma=uma3):
    c.append(maj, [cin, b[0], a[0]])
    for i in range(1, len(b)):
        c.append(maj, [a[i-1], b[i], a[i]])

    c.cx(a[-1], cout)

    for i in reversed(range(1, len(b))):
        c.append(uma, [a[i-1], b[i], a[i]])
    c.append(uma, [cin, b[0], a[0]])

def generate_adder_circuit(n, uma=uma3):
    if n % 2 != 0:
        raise ValueError('Odd number of qubits')

    cin = QuantumRegister(1, 'c_in')
    a = QuantumRegister(n//2-1, 'a')
    b = QuantumRegister(n//2-1, 'b->s')
    cout = QuantumRegister(1, 'c_out')
    c = QuantumCircuit(cin, b, a, cout)
        
    cuccaro_adder(c, cin, a, b, cout, uma=uma)
    return c

print('Cuccaro Adder:')
cuccaro8 = generate_adder_circuit(8)
cuccaro8.draw(fold=-1)
```

Output:
```
MAJ:
          ┌───┐     
q_0: ─────┤ X ├──■──
     ┌───┐└─┬─┘  │  
q_1: ┤ X ├──┼────■──
     └─┬─┘  │  ┌─┴─┐
q_2: ──■────■──┤ X ├
               └───┘
UMA2:
          ┌───┐     
q_0: ──■──┤ X ├──■──
       │  └─┬─┘┌─┴─┐
q_1: ──■────┼──┤ X ├
     ┌─┴─┐  │  └───┘
q_2: ┤ X ├──■───────
     └───┘          
UMA3:
                         ┌───┐     
q_0: ───────■────■───────┤ X ├─────
     ┌───┐┌─┴─┐  │  ┌───┐└─┬─┘┌───┐
q_1: ┤ X ├┤ X ├──■──┤ X ├──┼──┤ X ├
     └───┘└───┘┌─┴─┐└───┘  │  └─┬─┘
q_2: ──────────┤ X ├───────■────■──
               └───┘               
Cuccaro Adder:
        ┌──────┐                                       ┌───────┐
  c_in: ┤0     ├───────────────────────────────────────┤0      ├
        │      │                                       │       │
b->s_0: ┤1     ├───────────────────────────────────────┤1      ├
        │      │┌──────┐                      ┌───────┐│       │
b->s_1: ┤  MAJ ├┤1     ├──────────────────────┤1      ├┤  UMA3 ├
        │      ││      │┌──────┐
```

```python
def count_cnots(circuit, ignore_toffoli=False):
    '''A helper function that may or may not be useful.'''
    cnot_count = 0
    for [gate, qubits, cbits] in circuit:
        if len(qubits) <= 1 or (ignore_toffoli and len(qubits) == 3):
            pass
        elif len(qubits) == 2:
            cnot_count += 1
        else:
            raise ValueError('The circuit contains large gates that have not been decomposed')
    return cnot_count
```

```python
# Export to QASM
print(cuccaro8.qasm())
```

Output:
```
OPENQASM 2.0;
include "qelib1.inc";
gate gate_MAJ q0,q1,q2 { cx q2,q1; cx q2,q0; ccx q0,q1,q2; }
gate gate_UMA3 q0,q1,q2 { x q1; cx q0,q1; ccx q0,q1,q2; x q1; cx q2,q0; cx q2,q1; }
qreg c_in[1];
qreg b->s[3];
qreg a[3];
qreg c_out[1];
gate_MAJ c_in[0],b->s[0],a[0];
gate_MAJ a[0],b->s[1],a[1];
gate_MAJ a[1],b->s[2],a[2];
cx a[2],c_out[0];
gate_UMA3 a[1],b->s[2],a[2];
gate_UMA3 a[0],b->s[1],a[1];
gate_UMA3 c_in[0],b->s[0],a[0];
```

```python
# Try it out!
n = 8
cin = QuantumRegister(1, 'c_in')
a = QuantumRegister(n//2-1, 'a')
b = QuantumRegister(n//2-1, 'b->s')
cout = QuantumRegister(1, 'c_out')
creg = ClassicalRegister(n//2, 'sum')
c = QuantumCircuit(cin, b, a, cout, creg)

# Test on a specific input
# carry-in=0 CHANGE ME!
###c.x(cin[0])

# a=2 (0b010) CHANGE ME!
###c.x(a[0])
c.x(a[1])
###c.x(a[2])

# b=3 (0b011) CHANGE ME!
c.x(b[0])
c.x(b[1])
###c.x(b[2])

#cuccaro_adder(c, cin, a, b, cout)
c.append(cuccaro8.to_gate(label='CUCC'), [*cin, *b, *a, *cout])

# Measure the result
c.barrier()
c.measure([*b, *cout], creg)

# Show the circuit
display(c.draw(fold=-1))

# Simulate the result
# 0 + 2 + 3 = 5 (0b0101)
execute(c, backend, shots=128).result().get_counts(c)
```

Output:
```
┌───────┐ ░             
  c_in: ─────┤0      ├─░─────────────
        ┌───┐│       │ ░ ┌─┐         
b->s_0: ┤ X ├┤1      ├─░─┤M├─────────
        ├───┤│       │ ░ └╥┘┌─┐      
b->s_1: ┤ X ├┤2      ├─░──╫─┤M├──────
        └───┘│       │ ░  ║ └╥┘┌─┐   
b->s_2: ─────┤3      ├─░──╫──╫─┤M├───
             │  CUCC │ ░  ║  ║ └╥┘   
   a_0: ─────┤4      ├─░──╫──╫──╫────
        ┌───┐│       │ ░  ║  ║  ║    
   a_1: ┤ X ├┤5      ├─░──╫──╫──╫────
        └───┘│       │ ░  ║  ║  ║    
   a_2: ─────┤6      ├─░──╫──╫──╫────
             │       │ ░  ║  ║  ║ ┌─┐
 c_out: ─────┤7      ├─░──╫──╫──╫─┤M├
             └───────┘ ░  ║  ║  ║ └╥┘
 sum: 4/══════════════════╩══╩══╩══╩═
                          0  1  2  3 
{'0101': 128}
```

# 1. Local Circuit Optimizations

Now that you have decomposed the circuit into one- and two-qubit gates, you will look at optimizing the circuit at the one- and two-qubit gate level.  To keep things simple, you will only optimize the 8-qubit Cuccaro adder shown below.

```python
cuccaro8 = generate_adder_circuit(8)

print('The fully decomposed Cuccaro adder (Qiskit default decomposition):')
cucc_decomp = qiskit.transpile(cuccaro8, basis_gates=['x', 'h', 't', 'tdg', 'cx'])
cucc_decomp.draw(fold=-1)
```

Output:
```
The fully decomposed Cuccaro adder (Qiskit default decomposition):
                       ┌───┐                                                  ┌───┐                                                                                                                                                                                                                                                                                                                                          ┌───┐      ┌───┐          
  c_in: ───────────────┤ X ├───────────────────■─────────────────────■────■───┤ T ├───■─────────■─────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────────■─────────────────────■────■───┤ T ├───■──┤ X ├──────────
        ┌───┐          └─┬─┘                   │     
```

## 1.1. Gate Commutation

Take a look at the decomposed Cuccaro adder above.  In this circuit of CNOT and single-qubit gates, there are several places where we can rearrange and cancel gates.  For this question, take this small piece of the circuit shown below.

Commute (i.e. move) the single-qubit gates as far right as possible in the circuit without changing the functionality.

```python
print('A small piece of the decomposed Cuccaro adder:')
subcirc = QuantumCircuit.from_qasm_str('OPENQASM 2.0;\ninclude "qelib1.inc";\nqreg q[2];\nt q[1];\ncx q[0],q[1];\nt q[0];\ntdg q[1];\ncx q[0],q[1];\nx q[1];\ncx q[0],q[1];\nt q[0];\n')
subcirc.draw(fold=-1)
```

Output:
```
A small piece of the decomposed Cuccaro adder:
                ┌───┐                ┌───┐
q_0: ───────■───┤ T ├───■─────────■──┤ T ├
     ┌───┐┌─┴─┐┌┴───┴┐┌─┴─┐┌───┐┌─┴─┐└───┘
q_1: ┤ T ├┤ X ├┤ Tdg ├┤ X ├┤ X ├┤ X ├─────
     └───┘└───┘└─────┘└───┘└───┘└───┘
```

<!-- grade_id: 1-1, points: 0, grade: false, solution: true -->
```python
optimized_1 = QuantumCircuit(2)

# Rearrange these gates for your solution
# Do not add or remove gates
# You should get a circuit with two CNOTs next to each other
# and two T gates next to each other
optimized_1.t(1)
optimized_1.cx(0, 1)
optimized_1.tdg(1)
optimized_1.cx(0, 1)
optimized_1.cx(0, 1)
optimized_1.x(1)
optimized_1.t(0)
optimized_1.t(0)


print('Your solution:')
print(repr(optimized_1.qasm()))
optimized_1.draw(fold=-1)
```

Output:
```
Your solution:
'OPENQASM 2.0;\ninclude "qelib1.inc";\nqreg q[2];\nt q[1];\ncx q[0],q[1];\ntdg q[1];\ncx q[0],q[1];\ncx q[0],q[1];\nx q[1];\nt q[0];\nt q[0];\n'
                                ┌───┐┌───┐
q_0: ───────■───────────■────■──┤ T ├┤ T ├
     ┌───┐┌─┴─┐┌─────┐┌─┴─┐┌─┴─┐├───┤└───┘
q_1: ┤ T ├┤ X ├┤ Tdg ├┤ X ├┤ X ├┤ X ├─────
     └───┘└───┘└─────┘└───┘└───┘└───┘
```

<!-- grade_id: 1-1-test, points: 20, grade: true, solution: false -->
```python
# Do not delete this cell
# The autograder will test the value stored in `optimized_1`.
from collections import Counter
assert len(optimized_1) == len(subcirc), (
    'Do not add or remove gates, only rearrange them')
assert (Counter((type(g), tuple(q)) for g, q, c in optimized_1)
        == Counter((type(g), tuple(q)) for g, q, c in subcirc)), (
    'Do not add or remove gates, only rearrange them')
assert '\ncx q[0],q[1];\ncx q[0],q[1];\n' in optimized_1.qasm(), (
    'The solution will contain two CNOT gates back-to-back '
    '(as a next step, you would normally remove them)')
assert len(optimized_1.qasm().split('\ncx')[-1].split('\nt q[0];')) == 3, (
    'The solution will contain two T gates back-to-back.  '
    'Remember to move them as far right in the circuit as possible')
# Hidden tests
print('PASS')
```

Output:
```
PASS
```

Now than you have commuted as many single-qubit gates as possible to the right, an optimizer can look for pairs of CNOT gates and remove them from the circuit.  An optimizer can also merge sequences of single qubit gates.  $TT=S$ so in this circuit an optimizer can replace the pair of T gates with an S gate.

## 1.2. KAK

In this question, you will use Qiskit to find the shortest sequence of gates needed to execute this subcircuit on a real device.  The code below uses Qiskit's implementation of the KAK matrix decomposition to optimally compile to one CNOT gate.  (Question continued below.)

```python
print('A small piece of the decomposed Cuccaro adder:')
subcirc = QuantumCircuit.from_qasm_str('OPENQASM 2.0;\ninclude "qelib1.inc";\nqreg q[2];\nt q[1];\ncx q[0],q[1];\nt q[0];\ntdg q[1];\ncx q[0],q[1];\nx q[1];\ncx q[0],q[1];\nt q[0];\n')
display(subcirc.draw(fold=-1))

# Calculate the unitary matrix for the subcircuit
u_subcirc = qiskit.quantum_info.Operator(subcirc).data
print('The unitary matrix for this subcircuit:')
print(u_subcirc.round(3))  # Round to simplify print output
```

Output:
```
A small piece of the decomposed Cuccaro adder:
                ┌───┐                ┌───┐
q_0: ───────■───┤ T ├───■─────────■──┤ T ├
     ┌───┐┌─┴─┐┌┴───┴┐┌─┴─┐┌───┐┌─┴─┐└───┘
q_1: ┤ T ├┤ X ├┤ Tdg ├┤ X ├┤ X ├┤ X ├─────
     └───┘└───┘└─────┘└───┘└───┘└───┘     
The unitary matrix for this subcircuit:
[[ 0.   +0.j     0.   +0.j     1.   +0.j     0.   +0.j   ]
 [ 0.   +0.j     0.707+0.707j  0.   +0.j     0.   +0.j   ]
 [ 1.   +0.j     0.   +0.j     0.   +0.j     0.   +0.j   ]
 [ 0.   +0.j     0.   +0.j     0.   +0.j    -0.707+0.707j]]
```

```python
# An object that can compute the KAK unitary matrix decomposition
# using the CNOT gate
# Source code: https://github.com/Qiskit/qiskit-terra/blob/6354f3f4fa7a294a0f8b0b54d6b5bc5cdae49eca/qiskit/quantum_info/synthesis/two_qubit_decompose.py#L635
kak = qiskit.quantum_info.synthesis.TwoQubitBasisDecomposer(
    qiskit.circuit.library.standard_gates.CXGate(), euler_basis='ZXZ')

# An object that represents the unitary in a cannonical form as three real numbers
# (ignoring single-qubit gates)
# These Weyl coordinates are (0.7853981633974483, 0.0, 0.0)
# The other matrices are extra single-qubit gates
weyl_coords = (qiskit.quantum_info.synthesis.two_qubit_decompose
               .TwoQubitWeylDecomposition(u_subcirc))
weyl_coords
```

Output:
```
TwoQubitWeylControlledEquiv.from_bytes(
    # TwoQubitWeylControlledEquiv(
    # 	global phase: 3π/2
    # 	     ┌─────────┐            ┌────────────┐┌──────────┐
    # 	q_0: ┤ Ry(π/2) ├────────────┤0           ├┤ Ry(-π/2) ├────────────────────────
    # 	     ├─────────┤┌──────────┐│  Rxx(-π/2) │├──────────┤┌──────────┐┌──────────┐
    # 	q_1: ┤ Rx(π/2) ├┤ Ry(3π/4) ├┤1           ├┤ Rx(-π/2) ├┤ Ry(3π/4) ├┤ Rx(-π/2) ├
    # 	     └─────────┘└──────────┘└────────────┘└──────────┘└──────────┘└──────────┘
    # )
    b'k05VTVBZAQB2AHsnZGVzY3InOiAnPGMxNicsICdmb3J0cmFuX29yZGVyJzogRmFsc2UsICdzaGFw'
    b'ZSc6ICg0LCA0KSwgfSAgICAgICAgICAgICAgICAgICAgICAgICAgICAgICAgICAgICAgICAgICAg'
    b'ICAgICAgICAgICAgIAoAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAP7//////+8/AAAA'
    b'AAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAMs7f2aeoOY/yzt/Zp6g5j8AAAAA'
    b'AAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAPA/AAAAAAAAAAAAAAAAAAAAAAAAAAAA'
    b'AAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAAA
```

```python
print('Number of CNOT gates needed for this KAK decomposition:')
cnots_required = kak.num_basis_gates(u_subcirc)
cnots_required
```

Output:
```
Number of CNOT gates needed for this KAK decomposition:
1
```

```python
# Only one CNOT is required so use `decomp1` to find the four
# single-qubit unitary matricies needed
u0l, u1l, u0r, u1r = kak.decomp1(weyl_coords)
print('Single qubit gates needed around the one CNOT gate:')
# Left on qubit 0, left on 1, right on 0, right on 1
u0l, u1l, u0r, u1r
```

Output:
```
Single qubit gates needed around the one CNOT gate:
(array([[ 6.12323400e-17+1.01465364e-17j, -2.75374332e-33-1.00000000e+00j],
        [ 1.15499891e-33-1.00000000e+00j,  6.12323400e-17+1.01465364e-17j]]),
 array([[ 0.65487293-0.27125725j, -0.26993724+0.65168614j],
        [ 0.26993724+0.65168614j,  0.65487293+0.27125725j]]),
 array([[-3.80714113e-16+2.12520397e-16j,  7.07106781e-01+7.07106781e-01j],
        [-7.07106781e-01+7.07106781e-01j, -6.39763793e-17+3.23542700e-16j]]),
 array([[ 0.65487293-0.27125725j, -0.26993724-0.65168614j],
        [ 0.26993724-0.65168614j,  0.65487293+0.27125725j]]))
```

```python
# Make a circuit with these computed single-qubit unitaries
print('Optimal circuit if CNOT gates')
optimized_cnot = QuantumCircuit(2)
optimized_cnot.unitary(u0l, 0)
optimized_cnot.unitary(u1l, 1)
optimized_cnot.cx(0, 1)
optimized_cnot.unitary(u0r, 0)
optimized_cnot.unitary(u1r, 1)
display(optimized_cnot.draw(fold=-1))

# Check that it is correct
u_optimized_cnot = qiskit.quantum_info.Operator(optimized_cnot).data
assert matrix_equal(u_optimized_cnot, u_subcirc,
    ignore_phase=True, atol=1e-8, rtol=1e-8), (
    'Circuits are not equivalent')
print('Circuits are equivalent!')
```

Output:
```
Optimal circuit if CNOT gates
     ┌─────────┐     ┌─────────┐
q_0: ┤ Unitary ├──■──┤ Unitary ├
     ├─────────┤┌─┴─┐├─────────┤
q_1: ┤ Unitary ├┤ X ├┤ Unitary ├
     └─────────┘└───┘└─────────┘
Circuits are equivalent!
```

```python
# Instead of all the above steps, ask Qiskit to do it automatically
optimized_cnot_auto = kak(u_subcirc)
optimized_cnot_auto.draw(fold=-1)
```

Output:
```
global phase: 5π/4
        ┌───────┐                      ┌──────────┐  ┌───────┐ 
q40_0: ─┤ Rx(π) ├──────────────────■───┤ Rz(-π/2) ├──┤ Rx(π) ├─
       ┌┴───────┴┐┌─────────────┐┌─┴─┐┌┴──────────┴┐┌┴───────┴┐
q40_1: ┤ Rz(π/4) ├┤ Rx(-1.5659) ├┤ X ├┤ Rx(1.5659) ├┤ Rz(π/4) ├
       └─────────┘└─────────────┘└───┘└────────────┘└─────────┘
```

#### (1.2 continued)

For this problem, use KAK to compile `subcirc` for a device that only supports the **iSWAP** and single-qubit gates.

<!-- grade_id: 1-2, points: 0, grade: false, solution: true -->
```python
from qiskit import QuantumCircuit
from qiskit.quantum_info import Operator, TwoQubitBasisDecomposer
from qiskit.circuit.library import iSwapGate

unitary = Operator(subcirc)

# Initialize the KAK decomposer for iSWAP
kak_decomposer = TwoQubitBasisDecomposer(iSwapGate())

# Apply the KAK decomposition
decomposed_circuit = kak_decomposer(unitary)

# Convert the decomposed circuit into a QuantumCircuit
optimized_circuit = QuantumCircuit(decomposed_circuit.num_qubits)
optimized_circuit.append(decomposed_circuit, range(decomposed_circuit.num_qubits))
optimized_iswap = optimized_circuit.decompose()

print('Your optimized circuit for an iSWAP-capable device:')
optimized_iswap.draw(fold=-1)
```

Output:
```
Your optimized circuit for an iSWAP-capable device:
global phase: 7π/4
     ┌──────────────────┐┌────────┐ ┌────────────────┐┌────────┐┌─────────────────┐
q_0: ┤ U3(π/2,-π/2,π/2) ├┤0       ├─┤ U3(π,π/4,-π/4) ├┤0       ├┤ U3(π/2,-π/2,-π) ├
     ├─────────────────┬┘│  Iswap │┌┴────────────────┤│  Iswap │├─────────────────┤
q_1: ┤ U3(π/2,0,-3π/4) ├─┤1       ├┤ U3(π/2,-π,-π/2) ├┤1       ├┤ U3(π/2,-3π/4,0) ├
     └─────────────────┘ └────────┘└─────────────────┘└────────┘└─────────────────┘
```

<!-- grade_id: 1-2-test, points: 10, grade: true, solution: false -->
```python
# Do not delete this cell
# The autograder will test the value stored in `optimized_iswap`.
from collections import Counter
assert len(optimized_iswap.qasm().split('\ncx')) == 1, (
    'The solution should not contain any CNOT gates, only iSWAP (and single-qubit)')
assert len(optimized_iswap.qasm().split('\niswap')) <= 3, (
    'The optimal solution will contain two iSWAP gates')
u_optimized_iswap = qiskit.quantum_info.Operator(optimized_iswap).data
assert matrix_equal(u_optimized_iswap, u_subcirc,
    ignore_phase=True, atol=1e-8, rtol=1e-8), (
    'The compiled and original circuits are not equivalent')
# Hidden tests
print('PASS')
```

Output:
```
PASS
```

# 2. Global Optimization with the ZX Calculus

Recently, many powerful tools for circuit optimization have been developed that use a tool called the ZX Calculus.  After converting a quantum circuit into a ZX diagram it is easier to find simplifications, using the rules of the ZX calculus, that lead to smaller circuits.  You can [read more about ZX calculus here](https://zxcalculus.com/intro.html).

Follow the steps in the code below.

```python
# Install pyzx on your computer
!pip install pyzx==0.6.4
```

Output:
```
Collecting pyzx==0.6.4
  Downloading pyzx-0.6.4-py3-none-any.whl (271 kB)
[2K     [90m━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━[0m [32m271.1/271.1 kB[0m [31m5.4 MB/s[0m eta [36m0:00:00[0ma [36m0:00:01[0m
[?25hRequirement already satisfied: typing-extensions>=3.7.4 in /Users/jacqueslethuaut/miniconda3/envs/venv_quantum_conda/lib/python3.8/site-packages (from pyzx==0.6.4) (4.9.0)
Requirement already satisfied: numpy>=1.14 in /Users/jacqueslethuaut/miniconda3/envs/venv_quantum_conda/lib/python3.8/site-packages (from pyzx==0.6.4) (1.24.4)
Collecting pyperclip>=1.8.1 (from pyzx==0.6.4)
  Downloading pyperclip-1.8.2.tar.gz (20 kB)
  Preparing metadata (setup.py) ... [?25ldone
[?25hBuilding wheels for collected packages: pyperclip
  Building wheel for pyperclip (setup.py) ... [?25ldone
[?25h  Created wheel for pyperclip: filename=pyperclip-1.8.2-py3-none-any.whl size=11125 sha256=ac8253a5ca4d75e2b6143246662dc3e1f0e408e77c9fefb8b1938b9ddc880393
  Stored in directory: /Users/jacqu
```

```python
import pyzx as zx

# The cuccaro adder circuit
cuccaro = generate_adder_circuit(20)
cucc_decomp = qiskit.transpile(cuccaro, basis_gates=['x', 'h', 't', 'tdg', 'cx'])

# Convert the circuit into a ZX diagram
g = zx.Circuit.from_qasm(cuccaro.qasm())
print('The ZX diagram of the decomposed cuccaro adder:')
zx.draw(g)
```

Output:
```
The ZX diagram of the decomposed cuccaro adder:
<IPython.core.display.HTML object>
```

```python
# Simplify the diagram by replacing yellow squares (Hadamards)
g = g.to_basic_gates().to_graph()
zx.draw(g)
```

Output:
```
<IPython.core.display.HTML object>
```

## 2.1 Optimizer Routines

Now that we have a ZX diagram of the Cuccaro adder circuit, optimize the ZX diagram using one of the optimization routines listed here: [https://pyzx.readthedocs.io/en/latest/api.html](https://pyzx.readthedocs.io/en/latest/api.html#list-of-simplifications).  Find the one that optimizes this circuit the best.  I.e. the optimizer that results in a circuit with the fewest two-qubit gates.

```python
print(g.to_json())
```

Output:
```
{"wire_vertices": {"b0": {"annotation": {"boundary": true, "coord": [0, 0], "input": true, "output": false}}, "b1": {"annotation": {"boundary": true, "coord": [0, -1], "input": true, "output": false}}, "b2": {"annotation": {"boundary": true, "coord": [0, -2], "input": true, "output": false}}, "b3": {"annotation": {"boundary": true, "coord": [0, -3], "input": true, "output": false}}, "b4": {"annotation": {"boundary": true, "coord": [0, -4], "input": true, "output": false}}, "b5": {"annotation": {"boundary": true, "coord": [0, -5], "input": true, "output": false}}, "b6": {"annotation": {"boundary": true, "coord": [0, -6], "input": true, "output": false}}, "b7": {"annotation": {"boundary": true, "coord": [0, -7], "input": true, "output": false}}, "b8": {"annotation": {"boundary": true, "coord": [0, -8], "input": true, "output": false}}, "b9": {"annotation": {"boundary": true, "coord": [0, -9], "input": true, "output": false}}, "b10": {"annotation": {"boundary": true, "coord": [0, -10], "i
```

<!-- grade_id: 2-1, points: 0, grade: false, solution: true -->
```python
import matplotlib.pyplot as plt
# Reduce the circuit into a the form of a graph
# This will no longer look like a quantum circuit
# Additional work is needed to recover a valid quantum circuit
g_opt = zx.Graph.from_json(g.to_json())  # Convert to zx.Graph before optimizing

opt_methods = [zx.bialg_simp, zx.spider_simp, zx.id_simp, 
               zx.phase_free_simp, zx.pivot_simp, zx.pivot_gadget_simp, 
               zx.pivot_boundary_simp, zx.gadget_simp, zx.lcomp_simp, zx.clifford_simp, 
               zx.tcount, zx.to_gh, zx.to_rg, 
               zx.full_reduce, zx.teleport_reduce, zx.reduce_scalar, 
               zx.supplementarity_simp]
optimized_results = {}

for method in opt_methods:
    try:
        print(f"Testing method: {method.__name__}")
        g_copy = g_opt.copy()
        method(g_copy)
        
        vertices = list(g_copy.vertices())
        edges = list(g_copy.edges())

        print(f"Number of vertices: {len(vertices)}")
        print(f"Number of edges: {len(edges)}")

        # Draw and inspect the optimized graph
        #zx.draw(g_copy)
        
        optimized_results[method.__name__] = g_copy.copy()

    except Exception as e:
        print(f"Error during optimization with {method.__name__}: {e}")


def count_two_qubit_gates(qiskit_circuit):
    two_qubit_gates = 0
    for instr, qargs, _ in qiskit_circuit.data:
        if len(qargs) == 2:
            two_qubit_gates += 1
    return two_qubit_gates

# Dictionary to store the count of two-qubit gates for each optimization method
two_qubit_gate_counts = {}

for method_name, optimized_graph in optimized_results.items():
    # Extract the PyZX circuit from the optimized graph
    optimized_circuit = zx.extract_circuit(optimized_graph)

    # Convert the PyZX circuit to QASM format
    qasm_str = optimized_circuit.to_qasm()

    try:
        qiskit_circuit = QuantumCircuit.from_qasm_str(qasm_str)
        print(f"Successfully converted {method_name} to Qiskit circuit.")
    except Exception as e:
        print(f"Error converting {method_name} to Qiskit circuit: {e}")

    # Count two-qubit gates
    gate_count = count_two_qubit_gates(qiskit_circuit)
    two_qubit_gate_counts[method_name] = gate_count

best_method = min(two_qubit_gate_counts, key=two_qubit_gate_counts.get)

print(f"Best optimization method: {best_method} with {two_qubit_gate_counts[best_method]} two-qubit gates")

g_opt = optimized_results[best_method].copy()
print(f"Final optimized graph: {g_opt}")


try:
    zx.draw(g_opt)
except Exception as e:
    print(f"Error drawing the graph: {e}")

g_opt = zx.Graph.from_json(g.to_json())  # Convert to zx.Graph before optimizing

zx.spider_simp(g_opt)

zx.draw(g_opt)
```

Output:
```
Testing method: bialg_simp
bialg_simp: 8.  1 iterations
Number of vertices: 512
Number of edges: 646
Testing method: spider_simp
spider_simp: 78. 50. 25. 13. 9. 9. 6.  7 iterations
Number of vertices: 338
Number of edges: 454
Testing method: id_simp
id_simp: 36.  1 iterations
Number of vertices: 492
Number of edges: 626
Testing method: phase_free_simp
spider_simp: 78. 50. 25. 13. 9. 9. 6.  7 iterations
Number of vertices: 338
Number of edges: 454
Testing method: pivot_simp
Number of vertices: 528
Number of edges: 662
Testing method: pivot_gadget_simp
Number of vertices: 528
Number of edges: 662
Testing method: pivot_boundary_simp
Number of vertices: 528
Number of edges: 662
Testing method: gadget_simp
Number of vertices: 528
Number of edges: 662
Testing method: lcomp_simp
Number of vertices: 528
Number of edges: 662
Testing method: clifford_simp
Number of vertices: 260
Number of edges: 466
Testing method: tcount
Number of vertices: 528
Number of edges: 662
Testing method: to_gh
Number 
```

```python
# Convert the graph back into a quantum circuit
# This may not find the smallest possible circuit
c_opt = zx.extract_circuit(g_opt.copy(), quiet=True, optimize_cnots=3)
zx.draw(c_opt)
```

Output:
```
<IPython.core.display.HTML object>
```

<!-- grade_id: 2-1-test, points: 10, grade: true, solution: false -->
```python
# Do not delete this cell
# The autograder will test the value stored in `c_opt`.
_optimized = QuantumCircuit.from_qasm_str(c_opt.to_qasm())
assert len(cucc_decomp.qubits) == 20, 'The wrong input circuit was optimized.'
assert len(_optimized.qubits) == 20, 'Wrong number of qubits in the output circuit'
# Hidden tests
```

```python
# Convert this circuit back into a Qiskit circuit
optimized = QuantumCircuit.from_qasm_str(c_opt.to_qasm())
print(f'Optimized circuit:')
optimized.draw(fold=-1)
```

Output:
```
Optimized circuit:
      ┌───┐┌───┐      ┌───┐                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                          
```

```python
print(f'Original circuit (for comparison):')
cucc_decomp.draw(fold=-1)
```

Output:
```
Original circuit (for comparison):
                                                     ┌───┐                                                  ┌───┐                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                    
```

## 2.2 Output Circuit

Inspect the above "optimized" quantum circuit.  Does it look similar in structure to the original circuit with some gates removed or changed or is the structure completely different?

<!-- grade_id: 2-2, points: 0, grade: false, solution: true -->
```python
answer = (  # 'similar' or 'different'
    'different'
)
```

<!-- grade_id: 2-2-test, points: 5, grade: true, solution: false -->
```python
# Do not delete this cell
# The autograder will test the value stored in `c_opt`.
assert answer in ('similar', 'different'), 'Invalid answer'
# Hidden tests
```